Ensure you are using Python 3.0 and above

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE". Delete the 'raise NotImplementedError()' once you write your code there

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

### Complete the following Python exercises and submit the notebook on class folder

```python
# Imports
import numpy as np
from nose.tools import *
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

### Part A - Python stuff

#### Question 1
Write a function that adds two numbers together. If the numbers result in decimals, make sure that they are rounded off to 2 significant decimals, i.e. 23.1234 + 12.4321 = 35.56

<!-- grade_id: cell-cb3cad34e6f02386, points: 0, grade: false, solution: true -->
```python
# warm up
def add_two_numbers(x,y):
    return round((x+y),2)
```

```python
#Run tests
assert_equal= add_two_numbers(15.5674, 15.5674)
print(assert_equal)
```

Output:
```
31.13
```

#### Question 2
Complete this function to provide the cummulative sum of the first N numbers. The numbers should be counted from 1,2,.... Tip: You may want to look at the range function of Python.
E.g. function(4) = 1+2+3+4 = 10

Add an additional check - if the number N provided is outside an acceptable range, return 0. The acceptable range is 0 - 200 inclusive (i.e. it includes 0 and 200)

<!-- grade_id: cell-ddc6a0cd82f6d7db, points: 0, grade: false, solution: true -->
```python
def sum_first_N_numbers(N):
    if 0 <= N <= 200:
        return sum(range(1, N + 1))
    else:
        return 0
```

```python
#Test 1
result1 = sum_first_N_numbers(4)
print("Result for N=4:", result1)
```

Output:
```
Result for N=4: 10
```

```python
#Test 2
result2 = sum_first_N_numbers(0)
print("Result for N=0:", result2)
```

Output:
```
Result for N=0: 0
```

```python
#Test 3
result3 = sum_first_N_numbers(300)
print("Result for N=300:", result3)
```

Output:
```
Result for N=300: 0
```

#### Question 3
Write a function to calculate the mean of a list of numbers. Numbers should be rounded off to 2 significant decimals if they are decimal numbers

E.g. function([2,2,2,2]) = 2

<!-- grade_id: cell-6a114e8a16020116, points: 0, grade: false, solution: true -->
```python
def get_mean(x):
    if not x:
        return None
    else:
        mean = sum(x) / len(x)
        rounded_mean = round(mean, 2)
        return rounded_mean
```

```python
# Tests 1
#assert_equal(get_mean([2,3,4,5]), 3.5)
numbers_list = [6, 7, 8, 9]
result = get_mean(numbers_list)
print(result)
```

Output:
```
7.5
```

#### Question 4
Complete the function to get the power of a number. Ensure that if it is a decimal number it is rounded off to 3 (THREE) decimal places. 

E.g. function(3,2) = 9

<!-- grade_id: cell-0d560c78a3ca2566, points: 0, grade: false, solution: true -->
```python
def get_i_power(x, i):
    result = x ** i

    if isinstance(result, float):
        result = round(result, 3)

    return result
```

```python
# Test 1
#assert_equal(get_i_power(2,1), 2)
result = get_i_power(2, 1)
print(result)
```

Output:
```
2
```

```python
# Tests 2
#assert_equal(get_i_power(3.1,5), 1286.292)
result2 = get_i_power(3.1, 5)
print(result2)
```

Output:
```
286.292
```

#### Question 5
Write a function to reverse a string.

E.g. function('Mountain') = 'niatnuoM'

<!-- grade_id: cell-afe0b2a267eede6f, points: 0, grade: false, solution: true -->
```python

def reverse_string(s):
    # Reverse string s
    # E.g. reverse_string('Mountain') = 'niatnuoM'
    raise NotImplementedError()
```

```python
def reverse_string(s):
    # Reverse string s
    # E.g. reverse_string('Mountain') = 'niatnuoM'
    reversed_s = s[::-1]
    return reversed_s


original_string = 'Mountain'
result = reverse_string(original_string)
print(result)



```

Output:
```
niatnuoM
```

#### Question 6
Given the assignment_dates list, can you return the days from this string

E.g. function(assignment_dates) = ['Friday', 'Saturday', 'Sunday']

<!-- grade_id: cell-63604151c8149a01, points: 0, grade: false, solution: true -->
```python
from datetime import datetime
def get_days(assignment_dates):
    day_names = []
    days_to_match = ['Friday', 'Saturday', 'Sunday']
    for date_str in assignment_dates:
        date_object = datetime.strptime(date_str, '%A %d %b %Y')
        day_name = date_object.strftime('%A')
        if day_name in days_to_match:
            day_names.append(day_name)

    return day_names
```

```python
assignment_dates = ['Friday 01 Jan 2022', 'Saturday 02 Jan 2022', 'Sunday 03 Jan 2022']

result = get_days(assignment_dates)
result

```

Output:
```
['Saturday', 'Sunday']
```

#### Question 7
Complete the get_day_numbers function to return the INTEGER values for the numbers in the date

E.g. get_day_numbers(['Tue 23 April 2017']) = [23]

```python
def get_day_numbers(date_list):
    return [int(date.split()[1]) for date in date_list]
```

```python
# Test1
assignment_dates = ['Friday 21 Jan 2017', 'Saturday 10 March 2016', 'Sunday 11 October 2015']
assert get_day_numbers(assignment_dates)

print(get_day_numbers(assignment_dates))
```

Output:
```
[21, 10, 11]
```

```python
# Test2
assignment_dates2 = ['Monday 29 Jan 2017', 'Tuesday 9 March 2016', 'Wednesday 22 October 2015']
assert get_day_numbers(assignment_dates2)

print(get_day_numbers(assignment_dates2))
```

Output:
```
[29, 9, 22]
```

```python

```

#### Question 8
Complete the function to get the year in the date. Make sure this stays a string not an Integer

<!-- grade_id: cell-d8fcc99e93e8e654, points: 0, grade: false, solution: true -->
```python
def get_years(date_list):
    return [date.split()[-1] for date in date_list]
```

```python
# Test1
assignment_dates = ['Friday 27 Jan 2017', 'Saturday 12 March 2016', 'Sunday 3 October 2015']
assert get_years(assignment_dates) 
print(get_years(assignment_dates))
```

Output:
```
['2017', '2016', '2015']
```

```python
# Test2
assignment_dates4 = ['Monday 17 Jan 2020', 'Tuesday 1 March 2021', 'Wednesday 30 October 2022']
assert get_years(assignment_dates4)
print(get_years(assignment_dates4))
```

Output:
```
['2020', '2021', '2022']
```

### Part B: Numpy

Consider the array below

```python
import numpy as np
```

```python
a = np.random.rand(4,5)
a
```

Output:
```
array([[0.04931579, 0.76877252, 0.48256964, 0.85699407, 0.8987068 ],
       [0.44684858, 0.36803323, 0.57358367, 0.93653926, 0.85013128],
       [0.11955491, 0.23244941, 0.26658582, 0.59220242, 0.34583759],
       [0.12431123, 0.82445909, 0.21949538, 0.70881886, 0.5722664 ]])
```

<!-- grade_id: cell-46900ace5babb8cf, points: 2, grade: true, solution: true -->
```python
import numpy as np
a = np.random.rand(4,5)
ans1 = a[:, 1:4]
print(ans1)
```

Output:
```
[[0.74334313 0.06025087 0.75118264]
 [0.37486649 0.41934094 0.98407558]
 [0.37989705 0.87222506 0.561719  ]
 [0.88685641 0.95477851 0.0575479 ]]
```

<!-- grade_id: cell-b0c5638d712fc18e, points: 2, grade: true, solution: true -->
```python
import numpy as np 

ans2 = a[1:3, 0:3]

# Print the result
print(ans2)
```

Output:
```
[[0.19887065 0.70115531 0.30161066]
 [0.75979967 0.90756489 0.16983541]]
```

<!-- grade_id: cell-8e3b647ae1383de4, points: 5, grade: true, solution: true -->
```python
import numpy as np
arows_mean = np.mean(a, axis=1)
arows_sd = np.std(a, axis=1)

# Print the results
print("Mean of the rows:", arows_mean)
print("Standard deviation of the rows:", arows_sd)
```

Output:
```
Mean of the rows: [0.67288967 0.35153664 0.44942238 0.30295799]
Standard deviation of the rows: [0.24194918 0.23153425 0.33958841 0.3497348 ]
```

<!-- grade_id: cell-b3181cc311873032, points: 5, grade: true, solution: true -->
```python
import numpy as np
acols_mean = np.mean(a, axis=0)
acols_var = np.var(a, axis=0)

# Print the results
print("Mean of the columns:", acols_mean)
print("Variance of the columns:", acols_var)
```

Output:
```
Mean of the columns: [0.6897818  0.60641053 0.21839061 0.25244559 0.45397983]
Variance of the columns: [0.08690306 0.10337701 0.00876816 0.03719336 0.1275858 ]
```

Use numpy to load up the datafile `irisdata.txt`. Also load up the labels that correspond to this file `irislabels.txt`

<!-- grade_id: cell-37c4167b1c67a374, points: 4, grade: true, solution: true -->
```python
import numpy as np

# Load the data from irisdata.txt
data_file_path = 'irisdata.txt' 
iris_data = np.loadtxt(data_file_path, delimiter=',')

# Load the labels from irislabels.txt
labels_file_path = 'irislabels.txt' 
iris_labels = np.loadtxt(labels_file_path, dtype=int)

# Print the loaded data and labels
print("Iris Data:\n", iris_data)
print("\nIris Labels:\n", iris_labels)
```

Output:
```
Iris Data:
 [[5.1 3.5 1.4 0.2]
 [4.9 3.  1.4 0.2]
 [4.7 3.2 1.3 0.2]
 [4.6 3.1 1.5 0.2]
 [5.  3.6 1.4 0.2]
 [5.4 3.9 1.7 0.4]
 [4.6 3.4 1.4 0.3]
 [5.  3.4 1.5 0.2]
 [4.4 2.9 1.4 0.2]
 [4.9 3.1 1.5 0.1]
 [5.4 3.7 1.5 0.2]
 [4.8 3.4 1.6 0.2]
 [4.8 3.  1.4 0.1]
 [4.3 3.  1.1 0.1]
 [5.8 4.  1.2 0.2]
 [5.7 4.4 1.5 0.4]
 [5.4 3.9 1.3 0.4]
 [5.1 3.5 1.4 0.3]
 [5.7 3.8 1.7 0.3]
 [5.1 3.8 1.5 0.3]
 [5.4 3.4 1.7 0.2]
 [5.1 3.7 1.5 0.4]
 [4.6 3.6 1.  0.2]
 [5.1 3.3 1.7 0.5]
 [4.8 3.4 1.9 0.2]
 [5.  3.  1.6 0.2]
 [5.  3.4 1.6 0.4]
 [5.2 3.5 1.5 0.2]
 [5.2 3.4 1.4 0.2]
 [4.7 3.2 1.6 0.2]
 [4.8 3.1 1.6 0.2]
 [5.4 3.4 1.5 0.4]
 [5.2 4.1 1.5 0.1]
 [5.5 4.2 1.4 0.2]
 [4.9 3.1 1.5 0.1]
 [5.  3.2 1.2 0.2]
 [5.5 3.5 1.3 0.2]
 [4.9 3.1 1.5 0.1]
 [4.4 3.  1.3 0.2]
 [5.1 3.4 1.5 0.2]
 [5.  3.5 1.3 0.3]
 [4.5 2.3 1.3 0.3]
 [4.4 3.2 1.3 0.2]
 [5.  3.5 1.6 0.6]
 [5.1 3.8 1.9 0.4]
 [4.8 3.  1.4 0.3]
 [5.1 3.8 1.6 0.2]
 [4.6 3.2 1.4 0.2]
 [5.3 3.7 1.5 0.2]
 [5.  3.3 1.4 0.2]
 [7.  3.2 4.7 1.4]
 [6.4 3.2 4.5 1.5]

```

1. What are the size of the data from the datafile? 
2. The irislabels file has the labels or classes corresponding to the data what are the sizes of the different classes ?

<!-- grade_id: cell-1a3ec6ca5d0c1ee2, points: 6, grade: true, solution: true -->
```python
data_size = iris_data.shape
print("Size of the data:", data_size)


unique_classes, class_counts = np.unique(iris_labels, return_counts=True)
class_sizes = dict(zip(unique_classes, class_counts))
print("\nSizes of different classes:")
for class_label, class_size in class_sizes.items():
    print(f"Class {class_label}: {class_size}")
```

Output:
```
Size of the data: (150, 4)

Sizes of different classes:
Class 0: 50
Class 1: 50
Class 2: 50
```

#### Manually building a prototype based classifier

A prototype classifier obtains a very good representative of each class - also called a prototype and uses this for the classification task by computing the distance between a test set and the prototypes of each class. We are going to try and build one using this data.

To start this we need to split our data into a training and test set. 

<!-- grade_id: cell-87db5de2af7ebdd0, points: 5, grade: true, solution: true -->
```python
import numpy as np
from sklearn.model_selection import train_test_split
```

<!-- grade_id: cell-d1a12d347ddd519c, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Now we want to get the prototypes. This process can be complicated but for now lets just get the average of each of the data examples representing each class. 

<!-- grade_id: cell-da768715c85fff5e, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b5036749ed4bcf76, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Now lets see how well your algorithm performs

<!-- grade_id: cell-55fcc77e8583f3d7, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b8ca262a964f59e3, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Part C: Pandas - some data analytics tasks

Load the dataset `uganda_produce.csv` using pandas `pandas.read_csv`

<!-- grade_id: cell-071d8da9c5fe0afe, points: 4, grade: true, solution: true -->
```python
#SOLUTION
```

```python
import pandas as pd
df = pd.read_csv("uganda_produce.csv")
```

**Questions**

1. How many unique items are there in the dataset (Hint: you may want to investigate the `unique` method of pandas)
3. Of these products [sweet potatoes, beans and cassava] which had the highest yield overall ? What about for the year 2006 alone ?
2. Which commodity(item) overall had the highest area harvested ?
4. Plot the trend of area harvested for bananas over the years

<!-- grade_id: cell-7f2ad888dd4f4214, points: 1, grade: true, solution: true -->
```python
unique_items_all_columns = df.stack().unique()

# Print the unique items across all columns
print("Unique items across all columns:")
print(unique_items_all_columns)
```

Output:
```
Unique items across all columns:
['Uganda' 'Bananas' 'Area Harvested' ... 92200.0 87100.0 80300.0]
```

<!-- grade_id: cell-627284a78066ae4f, points: 2, grade: true, solution: true -->
```python
# Display the first few rows of the dataset to understand its structure
print(df.head())

max_production_row = df.loc[df['Value'].idxmax()]

# Print information about the row with the highest production
print("Row with the highest production:")
print(max_production_row)
```

Output:
```
country     item         element Unit                  Date     Value
0  Uganda  Bananas  Area Harvested   Ha  1/1/2003 12:00:00 AM  140323.0
1  Uganda  Bananas  Area Harvested   Ha  1/1/2004 12:00:00 AM  141083.0
2  Uganda  Bananas  Area Harvested   Ha  1/1/2005 12:00:00 AM  141505.0
3  Uganda  Bananas  Area Harvested   Ha  1/1/2006 12:00:00 AM  141674.0
4  Uganda  Bananas  Area Harvested   Ha  1/1/2007 12:00:00 AM  141758.0
Row with the highest production:
country                               Uganda
item       Fruit excl Melons,Total + (Total)
element                  Production Quantity
Unit                                  tonnes
Date                    1/1/2003 12:00:00 AM
Value                             10357475.0
Name: 528, dtype: object
```

<!-- grade_id: cell-7d6a7ea4969781d7, points: 2, grade: true, solution: true -->
```python
#THE-END
```

```python

```

<!-- grade_id: cell-84618533f475324e, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-34d1aba5d7cc55d8, points: 2, grade: true, solution: false -->
```python
# Tests 1
assert_equal(sum_first_N_numbers(4), 10)
```

<!-- grade_id: cell-38e9dccd85269c70, points: 3, grade: true, solution: false -->
```python
# Run tests
assert_equal(add_two_numbers(34.257, 42.936), 77.19)
```

<!-- grade_id: cell-a9ade0586e328717, points: 3, grade: true, solution: false -->
```python
# Tests 3
assert_equal(sum_first_N_numbers(300), 0)
```

<!-- grade_id: cell-93e0ff65d1eab080, points: 2, grade: true, solution: false -->
```python
# Tests 2
assert_equal(sum_first_N_numbers(0), 0)
```

<!-- grade_id: cell-bb8d92f5fdb244ec, points: 2, grade: true, solution: false -->
```python
# Tests 1
assert_equal(get_mean([2,3,4,5]), 3.5)
```

<!-- grade_id: cell-154da60daea46d99, points: 2, grade: true, solution: false -->
```python
# Test 1
assert_equal(get_i_power(3,2), 9)
```

<!-- grade_id: cell-52043c6a51a27e2c, points: 3, grade: true, solution: false -->
```python
# Tests 2
assert_equal(get_i_power(2.6,5), 118.814)
```

<!-- grade_id: cell-01861b730dc74d86, points: 4, grade: true, solution: false -->
```python
# Test1
assignment_dates = ['Friday 27 Jan 2017', 'Saturday 12 March 2016', 'Sunday 3 October 2015']
assert_equal(get_day_numbers(assignment_dates), [27, 12, 3])
```

<!-- grade_id: cell-9899e00775c1dc99, points: 0, grade: false, solution: true -->
```python
def get_day_numbers(date_list):
    return [int(date.split()[1]) for date in date_list]
```

<!-- grade_id: cell-826206bbf3aa056f, points: 3, grade: true, solution: false -->
```python
# Test2
assignment_dates2 = ['Monday 17 Jan 2017', 'Tuesday 1 March 2016', 'Wednesday 30 October 2015']
assert_equal(get_day_numbers(assignment_dates2), [17, 1, 30])
```

<!-- grade_id: cell-5f711d8c81943397, points: 3, grade: true, solution: false -->
```python
# Test1
assignment_dates = ['Friday 27 Jan 2017', 'Saturday 12 March 2016', 'Sunday 3 October 2015']
assert_equal(get_years(assignment_dates), ['2017', '2016', '2015'])
```

<!-- grade_id: cell-6f005e652741f3ff, points: 3, grade: true, solution: false -->
```python
# Test2
assignment_dates4 = ['Monday 17 Jan 2000', 'Tuesday 1 March 1998', 'Wednesday 30 October 2043']
assert_equal(get_years(assignment_dates4), ['2000', '1998', '2043'])
```

<!-- grade_id: cell-4551b5f0d1528d7c, points: 3, grade: true, solution: false -->
```python
# Test 1
assignment_dates = ['Friday 27 Jan 2017', 'Saturday 12 March 2016', 'Sunday 3 October 2015']
assert_equal(get_days(assignment_dates), ['Friday', 'Saturday', 'Sunday'])
```

<!-- grade_id: cell-d83aee453ba53d9f, points: 3, grade: true, solution: false -->
```python
assignment_dates2 = ['Monday 17 Jan 2017', 'Tuesday 1 March 2016', 'Wednesday 30 October 2015']
assert_equal(get_days(assignment_dates2), ['Monday', 'Tuesday', 'Wednesday'])
```

<!-- grade_id: cell-2bc2d51ae2e141ca, points: 2, grade: true, solution: false -->
```python
# Tests 3
assignment_dates3 = ['Thursday 27 Jan 2017']
assert_equal(get_days(assignment_dates3), ['Thursday'])
```

<!-- grade_id: cell-dfceb8f1cbacebb3, points: 3, grade: true, solution: false -->
```python
# Test 1
assert_equal(reverse_string('Mount Meru'), 'ureM tnuoM')
```